<a href="https://colab.research.google.com/github/aliraza33833/CE880labcoursework/blob/main/week05/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ali Raza"
Registration number(s) = "2111909"
```

---

---
# Welcome to CE880
### This is your week-5 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 

```python
# install sklearn package if you are using local computer and if you are using Google Coalb then skip this step
!pip install -U scikit-learn
import sys
import warnings

if not sys.warnoptions:
    warnings.simplefilter("ignore")
```

---
## Question 1: 
We’re ready to split a larger dataset to solve a regression problem. You’ll use a well-known Boston house prices dataset, which is included in sklearn. This dataset has 506 samples, 13 input variables, and the house values as the output. We have loaded dataset for you and splited into training and test set. Write code to fit linear regression on training `input (x)` and training `output (y)` and print the training score using `model.score` function. 

```python
from sklearn.datasets import load_boston
from sklearn.model_selection import train_test_split
x, y = load_boston(return_X_y=True)
x_train, x_test, y_train, y_test = train_test_split(x, y, test_size=0.4, random_state=0)
```

<!-- grade_id: my_regression_ans, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LinearRegression

def my_regression(x_train, y_train):
    """Write code to fit linear regression on training input (x) and training output (y) and 
    print the training score using model.score function. """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: my_regression_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_regression(x_train, y_train), 0.7668, rel_tol = 0.05)
```

---
## Question 2: 

Now, implement `RandomForestRegressor` on the same dataset and check the performance on the `test` dataset. Hint: Please read the question carefuly. 

<!-- grade_id: my_RMregression_ans, points: 0, grade: false, solution: true -->
```python
from sklearn.ensemble import RandomForestRegressor

def my_RMregression(x_train, y_train, x_test, y_test):
    """Write code to fit RandomForestRegressor on training input (x) and training output (y) and 
    print the test score using model.score function. """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: my_RMregression_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_RMregression(x_train, y_train, x_test, y_test), 0.8325, rel_tol = 0.05)
```

---
## Question 3: 

Write code to fit classificaiton model LogisticRegression on training `input (x)` and training `output (y)` and 
print the predicted (x) that is your same data on which you have trained your model 

<!-- grade_id: my_classification_ans, points: 0, grade: false, solution: true -->
```python

import numpy as np
from sklearn.linear_model import LogisticRegression
from sklearn.metrics import classification_report, confusion_matrix
x = np.arange(10).reshape(-1, 1)
y = np.array([0, 0, 0, 0, 1, 1, 1, 1, 1, 1])

def my_classification(x, y):
    """Write code to fit classificaiton model LogisticRegression on training input (x) and training output (y) and 
    print the predicted (x) that is your same data on which you have trained your model """
    # YOUR CODE HERE
    raise NotImplementedError()


```

<!-- grade_id: my_classification_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
y_hat = my_classification(x, y)
y_hat = y_hat.tolist()
assert y_hat == [0, 0, 0, 1, 1, 1, 1, 1, 1, 1]
```

```python

```