<a href="https://colab.research.google.com/github/aliraza33833/CE880labcoursework/blob/main/week06/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ali Raza"
Registration number(s) = "2111909"
```

---

---
# Welcome to CE880
### This is your week-6 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 

```python
# install sklearn package if you are using local computer and if you are using Google Coalb then skip this step
!pip install -U scikit-learn

!pip install -U pingouin

import sys
import warnings
import numpy as np

if not sys.warnoptions:
    warnings.simplefilter("ignore")
```

---
## Question 1: 

#### Student’s t-test
Tests whether the means of two independent samples are significantly different and return p-value. Please consider $\alpha = 0.05$ 

Assumptions

* Observations in each sample are independent and identically distributed (iid).
* Observations in each sample are normally distributed.
* Observations in each sample have the same variance.

Interpretation

* H0: the means of the samples are equal.
* H1: the means of the samples are unequal.

<!-- grade_id: my_ttest_ans, points: 0, grade: false, solution: true -->
```python
from scipy.stats import ttest_ind
mu, sigma = 0, 1 # mean and standard deviation
data1 = np.random.normal(mu, sigma, 100)

mu, sigma = 3, 1 # mean and standard deviation
data2 = np.random.normal(mu, sigma, 100)


def my_ttest(data1, data2):
    """Write a code to tests whether the means of two independent samples are significantly
    different and return p-value. Please consider 𝛼 = 0.05  """
    stat, p = ttest_ind(data1, data2)
    print('Statistics=%.3f, p=%.3f' % (stat, p))
    # interpret
    alpha = 0.05
    if p > alpha:
        print('Same distributions (fail to reject H0)')
    else:
        print('Different distributions (reject H0)')
    # YOUR CODE HERE
    #raise NotImplementedError()
    return p
```

<!-- grade_id: my_ttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
val = my_ttest(data1, data2)
assert math.isclose(val, 0, abs_tol = 0.05)
```

Output:
```
Statistics=-19.038, p=0.000
Different distributions (reject H0)
```

---
## Question 2: 

#### Paired Student’s t-test
Tests whether the means of two paired samples are significantly different. Please consider $\alpha = 0.05$ 

Assumptions

* Observations in each sample are independent and identically distributed (iid).
* Observations in each sample are normally distributed.
* Observations in each sample have the same variance.
* Observations across each sample are paired.

Interpretation

* H0: the means of the samples are equal.
* H1: the means of the samples are unequal.

<!-- grade_id: my_pairedttest_ans, points: 0, grade: false, solution: true -->
```python
# Example of the Paired Student's t-test
from scipy.stats import ttest_rel
import pingouin as pt

data1 = [0.873, 2.817, 0.121, -0.945, -0.055, -1.436, 0.360, -1.478, -1.637, -1.869]
data2 = [1.142, -0.432, -0.938, -0.729, -0.846, -0.157, 0.500, 1.183, -1.075, -0.169]


def my_pairedttest(data1, data2):
    """Write code to tests whether the means of two paired samples are 
    significantly different. Please consider 𝛼=0.05 """
    stat, p = ttest_rel(data1, data2)
    test = pt.ttest(data1, data2, paired=True)
    return test.power
```

<!-- grade_id: my_pairedttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_pairedttest(data1, data2), 0.746, rel_tol = 0.05)
```

---
## Question 3: 

#### Shapiro-Wilk Test
Tests whether a data sample has a Gaussian distribution.

Assumptions

* Observations in each sample are independent and identically distributed (iid).

Interpretation

* H0: the sample has a Gaussian distribution.
* H1: the sample does not have a Gaussian distribution.

<!-- grade_id: my_normalityttest_ans, points: 0, grade: false, solution: true -->
```python
from scipy.stats import shapiro
data = [0.873, 2.817, 0.121, -0.945, -0.055, -1.436, 0.360, -1.478, -1.637, -1.869]

def my_normalityttest(data):
    """Write code to to check the normality of the data. In other words, 
    you are checking if the data is normally distributed and return p value. 
    Hint: Please see what is shapiro test and how you can use it
    https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.shapiro.html"""
    # YOUR CODE HERE
    shapiro_test = shapiro(data)
    pval = shapiro_test.pvalue
    return pval
my_normalityttest(data)
```

Output:
```
0.19340917468070984
```

<!-- grade_id: my_normalityttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_normalityttest(data), 0.193, rel_tol = 0.05)
```

```python

```